```python
# ── COLAB SETUP ────────────────────────────────────────────────────────────────
# Run this cell first to confirm your environment is ready.
#
# This notebook uses only Python standard library modules (math, os, etc.)
# and does not require any additional package installations.
# It runs identically in Google Colab, JupyterHub, or a local Jupyter environment.

import sys, platform
print(f"Python {sys.version}")
print(f"Platform: {platform.system()}")
print("Environment ready — no additional packages required ✓")

```

# Python for Chemists — Part 1: Core Concepts

This notebook introduces the Python skills you will use throughout your chemistry research.
Every concept is illustrated with chemical examples you already know from lecture — no prior
programming experience is assumed.

### Topics covered

| Section | Python concept | Chemistry application |
|---------|---------------|----------------------|
| 1 | Variables & data types | Storing physical quantities |
| 2 | Arithmetic & unit conversion | Stoichiometry calculations |
| 3 | Strings | Chemical formulas, element names |
| 4 | Lists | Spectral data, titration series |
| 5 | Loops | Processing datasets |
| 6 | Dictionaries | Periodic table look-ups |
| 7 | Functions | Reusable equations |
| 8 | File input / output | Reading experimental data |

> **How to use this notebook:** Read each explanation, run the **Example** cell with
> **Shift + Enter**, then complete the **✏️ Practice** cell on your own.
> Each practice cell has **variable stubs already named for you** — replace each `None`
> with the correct value, then delete the `raise NotImplementedError()` line.
> Run the **Test** cell to check your work: ✓ means you passed.


---
## Section 1 — Variables and Data Types

A **variable** is a named container for a value. You create one with a single `=` sign.
Python automatically determines the **type** of the value stored.

| Type | Python name | Typical chemistry use |
|------|-------------|----------------------|
| Whole number | `int` | Atomic number, number of moles (integer) |
| Decimal number | `float` | Molar mass, concentration, temperature |
| Text | `str` | Element symbol, compound name |
| True / False | `bool` | "Is the solution acidic?", convergence flags |

Good variable names are lowercase, descriptive, and use underscores for spaces:
`molar_mass_NaCl`, not `x` or `MolarMassNaCl`.


```python
# ── Example 1 ────────────────────────────────────────────────────────
atomic_number_Na  = 11           # int   — atomic number of sodium
molar_mass_NaCl   = 58.44        # float — g/mol
compound_name     = "sodium chloride"  # str
is_ionic          = True         # bool

print(f"Compound:      {compound_name}")
print(f"Atomic # (Na): {atomic_number_Na}")
print(f"Molar mass:    {molar_mass_NaCl:.2f} g/mol")
print(f"Ionic?         {is_ionic}")
print()
print(f"type(atomic_number_Na) → {type(atomic_number_Na)}")
print(f"type(molar_mass_NaCl)  → {type(molar_mass_NaCl)}")
print(f"type(compound_name)    → {type(compound_name)}")
print(f"type(is_ionic)         → {type(is_ionic)}")

```

#### ✏️ Practice 1
Create variables for the following properties of **caffeine (C₈H₁₀N₄O₂)**.
The variable names are already written for you — replace each `None` with the correct value.

- `molar_mass_caffeine`: 194.19 g/mol → should be a **float**
- `n_carbons`: 8 carbon atoms → should be an **int**
- `formula_caffeine`: molecular formula → should be a **str**
- `solid_at_RT`: is caffeine a solid at room temperature? → should be a **bool** (`True` or `False`, not `"yes"`)


<!-- grade_id: practice1_solution, points: 0, grade: false, solution: true -->
```python
# ── Practice 1 ───────────────────────────────────────────────────────
molar_mass_caffeine = None   # float  — molar mass in g/mol
n_carbons           = None   # int    — number of carbon atoms
formula_caffeine    = None   # str    — molecular formula
solid_at_RT         = None   # bool   — True or False (not "yes" / "no")

### BEGIN SOLUTION
molar_mass_caffeine = 194.19
n_carbons           = 8
formula_caffeine    = "C8H10N4O2"
solid_at_RT         = True
### END SOLUTION

print(f"Molar mass:   {molar_mass_caffeine} g/mol")
print(f"Carbon atoms: {n_carbons}")
print(f"Formula:      {formula_caffeine}")
print(f"Solid at RT:  {solid_at_RT}")

```

<!-- grade_id: practice1_tests, points: 4, grade: true, solution: false -->
```python
# ── Tests for Practice 1 ─────────────────────────────────────────────
assert isinstance(molar_mass_caffeine, float),     f"molar_mass_caffeine should be a float like 194.19, not {type(molar_mass_caffeine).__name__}"
assert abs(molar_mass_caffeine - 194.19) < 0.01,     f"Expected molar_mass_caffeine ≈ 194.19, got {molar_mass_caffeine}"

assert isinstance(n_carbons, int),     f"n_carbons should be an int, not {type(n_carbons).__name__}"
assert n_carbons == 8,     f"Expected n_carbons = 8, got {n_carbons}"

assert isinstance(formula_caffeine, str),     "formula_caffeine should be a string"
assert formula_caffeine == "C8H10N4O2",     f"Expected 'C8H10N4O2', got {formula_caffeine!r}"

assert isinstance(solid_at_RT, bool),     f"solid_at_RT must be True or False (a bool), not {type(solid_at_RT).__name__!r}. Hint: don't use quotes."
assert solid_at_RT is True,     "Caffeine is indeed a solid at room temperature — solid_at_RT should be True"

print("✓ All Practice 1 checks passed!")

```

---
## Section 2 — Arithmetic and Unit Conversion

| Operator | Meaning | Example |
|----------|---------|---------|
| `+`, `-`, `*`, `/` | add, subtract, multiply, divide | `M * n` |
| `**` | exponentiation | `10**(-7)` for 10⁻⁷ |
| `//` | integer (floor) division | `7 // 2 → 3` |
| `%` | modulus (remainder) | `7 % 2 → 1` |


```python
# ── Example 2 ────────────────────────────────────────────────────────
# How many grams of NaCl to make 250 mL of 0.500 M solution?
molarity_NaCl   = 0.500    # mol/L
volume_L        = 0.250    # L
molar_mass_NaCl = 58.44    # g/mol

moles_NaCl = molarity_NaCl * volume_L
mass_NaCl  = moles_NaCl * molar_mass_NaCl

print(f"Moles of NaCl needed: {moles_NaCl:.4f} mol")
print(f"Mass of NaCl needed:  {mass_NaCl:.2f} g")
print()

# Beer–Lambert law: A = ε × l × c
epsilon    = 5500.0   # L mol⁻¹ cm⁻¹
path_len   = 1.0      # cm
conc       = 2.5e-4   # mol/L
absorbance = epsilon * path_len * conc
print(f"Absorbance A = {absorbance:.4f}")

```

#### ✏️ Practice 2
A student dissolves **5.85 g of NaCl** in enough water to make **500 mL** of solution.
The given constants are already defined. Calculate the three result variables.

- `moles`: moles of NaCl = mass / molar mass
- `molarity`: molarity = moles / volume (in L)
- `A`: absorbance via Beer–Lambert (`A = ε·l·c`, ε = 5500 L mol⁻¹ cm⁻¹, l = 1 cm)


<!-- grade_id: practice2_solution, points: 0, grade: false, solution: true -->
```python
# ── Practice 2 ───────────────────────────────────────────────────────
# Given — do not change these:
mass_NaCl   = 5.85     # g
molar_mass  = 58.44    # g/mol
volume_L    = 0.500    # L
epsilon     = 5500.0   # L mol⁻¹ cm⁻¹
path_length = 1.0      # cm

# Calculate — replace each None:
moles    = None   # mol
molarity = None   # mol/L
A        = None   # absorbance (dimensionless)

### BEGIN SOLUTION
moles    = mass_NaCl / molar_mass
molarity = moles / volume_L
A        = epsilon * path_length * molarity
### END SOLUTION

print(f"Moles:      {moles:.5f} mol")
print(f"Molarity:   {molarity:.5f} mol/L")
print(f"Absorbance: {A:.4f}")

```

<!-- grade_id: practice2_tests, points: 4, grade: true, solution: false -->
```python
# ── Tests for Practice 2 ─────────────────────────────────────────────
assert abs(moles - (5.85 / 58.44)) < 1e-4,     f"moles should be ≈ {5.85/58.44:.5f}, got {moles:.5f}"

expected_M = (5.85 / 58.44) / 0.500
assert abs(molarity - expected_M) < 1e-4,     f"molarity should be ≈ {expected_M:.5f} mol/L, got {molarity:.5f}"

expected_A = 5500.0 * 1.0 * expected_M
assert abs(A - expected_A) < 0.01,     f"Absorbance should be ≈ {expected_A:.4f}, got {A:.4f}"

print("✓ All Practice 2 checks passed!")

```

---
## Section 3 — Strings

| Operation | What it does | Example |
|-----------|-------------|---------|
| `s.strip()` | remove leading/trailing whitespace | `"  H2O  ".strip()` → `"H2O"` |
| `s.split(sep)` | split on separator → list | `"H,C,N".split(",")` → `["H","C","N"]` |
| `s.replace(old, new)` | substitute text | |
| `sub in s` | membership test | `"H" in "H2O"` → `True` |
| `s.capitalize()` | capitalize first letter | `"acetone".capitalize()` → `"Acetone"` |


```python
# ── Example 3 ────────────────────────────────────────────────────────
formula = "  C6H12O6  "   # glucose, with accidental whitespace
clean = formula.strip()
print(f"Cleaned: '{clean}'  |  Contains oxygen? {'O' in clean}")

data_line = "glucose, 180.16, 6, solid"
name, molar_mass, n_carbons, phase = data_line.split(", ")
molar_mass_float = float(molar_mass)

label = f"{name.capitalize()} (M = {molar_mass_float:.1f} g/mol)"
print(f"Plot label: '{label}'")

```

#### ✏️ Practice 3
Parse this instrument export line into its parts.
The variable names are pre-written — replace each `None`.

```
"  acetic acid | 60.052 | 4 | liquid  "
```


<!-- grade_id: practice3_solution, points: 0, grade: false, solution: true -->
```python
# ── Practice 3 ───────────────────────────────────────────────────────
raw_line = "  acetic acid | 60.052 | 4 | liquid  "

cleaned          = None   # str — raw_line with whitespace removed
name             = None   # str — first field after splitting on " | "
molar_mass       = None   # str — second field
n_carbons        = None   # str — third field
phase            = None   # str — fourth field
molar_mass_float = None   # float — molar_mass converted to a number

### BEGIN SOLUTION
cleaned = raw_line.strip()
name, molar_mass, n_carbons, phase = cleaned.split(" | ")
molar_mass_float = float(molar_mass)
### END SOLUTION

print(f"{name.capitalize()} has a molar mass of {molar_mass_float:.2f} g/mol and is a {phase}.")

```

<!-- grade_id: practice3_tests, points: 4, grade: true, solution: false -->
```python
# ── Tests for Practice 3 ─────────────────────────────────────────────
assert cleaned == "acetic acid | 60.052 | 4 | liquid",     f"cleaned should have no leading/trailing spaces, got {cleaned!r}"
assert name == "acetic acid",     f"name should be 'acetic acid', got {name!r}"
assert phase == "liquid",     f"phase should be 'liquid', got {phase!r}"
assert isinstance(molar_mass_float, float),     f"molar_mass_float should be a float, got {type(molar_mass_float).__name__}"
assert abs(molar_mass_float - 60.052) < 1e-3,     f"molar_mass_float should be ≈ 60.052, got {molar_mass_float}"

print("✓ All Practice 3 checks passed!")

```

---
## Section 4 — Lists

| Operation | Example | Result |
|-----------|---------|--------|
| Index | `data[0]` | first item |
| Negative index | `data[-1]` | last item |
| Append | `data.append(x)` | add to end |
| Length | `len(data)` | number of items |
| Aggregates | `sum()`, `min()`, `max()` | built-in aggregates |


```python
# ── Example 4 ────────────────────────────────────────────────────────
absorbances = [0.082, 0.155, 0.241, 0.318, 0.397, 0.461]

print(f"First: {absorbances[0]}   Last: {absorbances[-1]}")
print(f"Max: {max(absorbances):.3f}   Min: {min(absorbances):.3f}")
print(f"Average: {sum(absorbances)/len(absorbances):.3f}")

absorbances.append(0.510)
print(f"After append: {absorbances}")

```

#### ✏️ Practice 4
pH values from an acid–base titration are given below.
All result variable names are pre-written — replace each `None`.

```
3.21, 3.45, 3.89, 4.62, 7.00, 9.38, 10.11, 10.54, 10.78
```


<!-- grade_id: practice4_solution, points: 0, grade: false, solution: true -->
```python
# ── Practice 4 ───────────────────────────────────────────────────────
ph_values      = None   # list  — all nine pH readings
initial_pH     = None   # float — first value in the list
final_pH       = None   # float — last value in the list
min_pH         = None   # float — lowest pH recorded
max_pH         = None   # float — highest pH recorded
equivalence_pH = None   # float — value at index 4 (midpoint of titration)
count_basic    = None   # int   — number of values greater than 7

### BEGIN SOLUTION
ph_values      = [3.21, 3.45, 3.89, 4.62, 7.00, 9.38, 10.11, 10.54, 10.78]
initial_pH     = ph_values[0]
final_pH       = ph_values[-1]
min_pH         = min(ph_values)
max_pH         = max(ph_values)
equivalence_pH = ph_values[4]
count_basic    = sum(1 for p in ph_values if p > 7)
### END SOLUTION

print(f"Initial pH:         {initial_pH}")
print(f"Final pH:           {final_pH}")
print(f"Min / Max:          {min_pH} / {max_pH}")
print(f"Equivalence point:  {equivalence_pH}")
print(f"Basic measurements: {count_basic}")

```

<!-- grade_id: practice4_tests, points: 4, grade: true, solution: false -->
```python
# ── Tests for Practice 4 ─────────────────────────────────────────────
assert ph_values == [3.21, 3.45, 3.89, 4.62, 7.00, 9.38, 10.11, 10.54, 10.78],     "ph_values does not match the expected list"
assert abs(initial_pH - 3.21) < 1e-6,    f"initial_pH should be 3.21, got {initial_pH}"
assert abs(final_pH   - 10.78) < 1e-6,   f"final_pH should be 10.78, got {final_pH}"
assert abs(min_pH     - 3.21) < 1e-6,    f"min_pH should be 3.21, got {min_pH}"
assert abs(max_pH     - 10.78) < 1e-6,   f"max_pH should be 10.78, got {max_pH}"
assert abs(equivalence_pH - 7.00) < 1e-6,     f"equivalence_pH (ph_values[4]) should be 7.00, got {equivalence_pH}"
assert count_basic == 4,     f"count_basic should be 4 (9.38, 10.11, 10.54, 10.78 are > 7), got {count_basic}"

print("✓ All Practice 4 checks passed!")

```

---
## Section 5 — Loops

```python
for item in collection:
    # do something with item
```

- `enumerate(collection)` — gives index and value at each step
- `zip(list1, list2)` — pairs items from two lists side by side


```python
# ── Example 5 ────────────────────────────────────────────────────────
absorbances = [0.082, 0.155, 0.241, 0.318, 0.397, 0.461]
epsilon, path_length = 5500.0, 1.0

concentrations = []
for i, A in enumerate(absorbances):
    time_min = i * 5
    c_uM = (A / (epsilon * path_length)) * 1e6
    concentrations.append(c_uM)
    print(f"t = {time_min:2d} min   A = {A:.3f}   c = {c_uM:.2f} μM")

```

#### ✏️ Practice 5
The two lists below give five solvents and their boiling points in °C.
Fill in the two result variables.

- `bp_kelvin`: a list of boiling points converted to Kelvin (`K = °C + 273.15`)
  — build it inside the loop using `zip`
- `highest_bp_solvent`: the name of the solvent with the highest boiling point


<!-- grade_id: practice5_solution, points: 0, grade: false, solution: true -->
```python
# ── Practice 5 ───────────────────────────────────────────────────────
# Given — do not change these:
solvents   = ["Diethyl ether", "Dichloromethane", "Acetone", "Ethanol", "Water"]
bp_celsius = [34.6, 39.6, 56.1, 78.4, 100.0]

# Results — fill these in:
bp_kelvin          = None   # list of floats — one Kelvin value per solvent
highest_bp_solvent = None   # str — name of the solvent with the highest bp

### BEGIN SOLUTION
bp_kelvin = []
for name, bp_C in zip(solvents, bp_celsius):
    bp_K = bp_C + 273.15
    bp_kelvin.append(bp_K)
    print(f"{name:<20}: {bp_K:.2f} K")

max_bp = max(bp_celsius)
highest_bp_solvent = solvents[bp_celsius.index(max_bp)]
### END SOLUTION

print(f"
Highest boiling point: {highest_bp_solvent}")

```

<!-- grade_id: practice5_tests, points: 4, grade: true, solution: false -->
```python
# ── Tests for Practice 5 ─────────────────────────────────────────────
expected_K = [bp + 273.15 for bp in [34.6, 39.6, 56.1, 78.4, 100.0]]
assert len(bp_kelvin) == 5, f"bp_kelvin should have 5 values, got {len(bp_kelvin)}"
for i, (got, exp) in enumerate(zip(bp_kelvin, expected_K)):
    assert abs(got - exp) < 1e-4,         f"bp_kelvin[{i}] should be {exp:.2f} K, got {got:.2f} K"
assert highest_bp_solvent == "Water",     f"highest_bp_solvent should be 'Water', got {highest_bp_solvent!r}"

print("✓ All Practice 5 checks passed!")

```

---
## Section 6 — Dictionaries

```python
d = {"key": value}
d["key"]          # retrieve
d["new_key"] = v  # add/update
d.items()         # all key-value pairs
```


```python
# ── Example 6 ────────────────────────────────────────────────────────
elements = {
    "H":  {"name": "Hydrogen",  "atomic_number": 1,  "atomic_mass": 1.008,  "electronegativity": 2.20},
    "C":  {"name": "Carbon",    "atomic_number": 6,  "atomic_mass": 12.011, "electronegativity": 2.55},
    "N":  {"name": "Nitrogen",  "atomic_number": 7,  "atomic_mass": 14.007, "electronegativity": 3.04},
    "O":  {"name": "Oxygen",    "atomic_number": 8,  "atomic_mass": 15.999, "electronegativity": 3.44},
    "F":  {"name": "Fluorine",  "atomic_number": 9,  "atomic_mass": 18.998, "electronegativity": 3.98},
    "Cl": {"name": "Chlorine",  "atomic_number": 17, "atomic_mass": 35.45,  "electronegativity": 3.16},
}

print(f"{'Symbol':<6} {'Name':<12} {'Mass (u)':>10} {'EN':>6}")
print("-" * 38)
for sym, p in elements.items():
    print(f"{sym:<6} {p['name']:<12} {p['atomic_mass']:>10.3f} {p['electronegativity']:>6.2f}")

```

#### ✏️ Practice 6
Using the `elements` dictionary above, compute the molar masses of two molecules
by looking up atomic masses from the dictionary — do **not** hard-code numbers.

- `M_H2O`: molar mass of H₂O = 2 × M(H) + 1 × M(O)
- `M_CH4`: molar mass of CH₄ = 1 × M(C) + 4 × M(H)


<!-- grade_id: practice6_solution, points: 0, grade: false, solution: true -->
```python
# ── Practice 6 ───────────────────────────────────────────────────────
M_H2O = None   # float — molar mass of water in g/mol
M_CH4 = None   # float — molar mass of methane in g/mol

### BEGIN SOLUTION
M_H2O = 2 * elements["H"]["atomic_mass"] + 1 * elements["O"]["atomic_mass"]
M_CH4 = 1 * elements["C"]["atomic_mass"] + 4 * elements["H"]["atomic_mass"]
### END SOLUTION

print(f"M(H₂O) = {M_H2O:.3f} g/mol")
print(f"M(CH₄) = {M_CH4:.3f} g/mol")

```

<!-- grade_id: practice6_tests, points: 4, grade: true, solution: false -->
```python
# ── Tests for Practice 6 ─────────────────────────────────────────────
expected_H2O = 2 * 1.008 + 15.999
expected_CH4 = 12.011 + 4 * 1.008

assert abs(M_H2O - expected_H2O) < 1e-3,     f"M_H2O should be ≈ {expected_H2O:.3f} g/mol, got {M_H2O:.3f}"
assert abs(M_CH4 - expected_CH4) < 1e-3,     f"M_CH4 should be ≈ {expected_CH4:.3f} g/mol, got {M_CH4:.3f}"

print("✓ All Practice 6 checks passed!")

```

---
## Section 7 — Functions

```python
def function_name(parameter1, parameter2=default):
    """Docstring."""
    result = ...
    return result
```


```python
# ── Example 7 ────────────────────────────────────────────────────────
import math

def ideal_gas_volume(n_moles, temp_C=25.0, pressure_atm=1.0):
    """Calculate volume (L) of an ideal gas: PV = nRT."""
    R = 0.082057   # L·atm / (mol·K)
    T = temp_C + 273.15
    return (n_moles * R * T) / pressure_atm

def pH_from_conc(H_concentration):
    """Calculate pH from [H⁺] in mol/L."""
    return -math.log10(H_concentration)

print(f"1 mol at 25°C, 1 atm: {ideal_gas_volume(1.0):.3f} L")
print(f"1 mol at  0°C, 1 atm: {ideal_gas_volume(1.0, temp_C=0.0):.3f} L  ← STP")
for H in [1e-2, 1e-7, 1e-12]:
    print(f"[H⁺] = {H:.0e}  →  pH = {pH_from_conc(H):.1f}")

```

#### ✏️ Practice 7
Two function shells are already written with the correct names, parameters, and docstrings.
Fill in the `return` statement for each one — replace the `pass`.

- `mass_needed(molarity, volume_mL, molar_mass)` → mass (g) = molarity × (volume_mL / 1000) × molar_mass
- `percent_yield(actual_g, theoretical_g)` → % yield = (actual / theoretical) × 100


<!-- grade_id: practice7_solution, points: 0, grade: false, solution: true -->
```python
# ── Practice 7 ───────────────────────────────────────────────────────
def mass_needed(molarity, volume_mL, molar_mass):
    """Return mass (g) of solute needed for a solution preparation."""
    ### BEGIN SOLUTION
    return molarity * (volume_mL / 1000) * molar_mass
    ### END SOLUTION

def percent_yield(actual_g, theoretical_g):
    """Return percent yield of a reaction."""
    ### BEGIN SOLUTION
    return (actual_g / theoretical_g) * 100
    ### END SOLUTION

# Test your functions:
mass = mass_needed(0.10, 250, 180.16)
pct  = percent_yield(3.87, 4.50)
print(f"Mass of glucose needed: {mass:.2f} g")
print(f"Percent yield:          {pct:.1f}%")

```

<!-- grade_id: practice7_tests, points: 4, grade: true, solution: false -->
```python
# ── Tests for Practice 7 ─────────────────────────────────────────────
assert callable(mass_needed),   "mass_needed should be a function"
assert callable(percent_yield), "percent_yield should be a function"

expected_mass = 0.10 * (250 / 1000) * 180.16
assert abs(mass_needed(0.10, 250, 180.16) - expected_mass) < 1e-3,     f"mass_needed(0.10, 250, 180.16) should be ≈ {expected_mass:.3f} g"
assert abs(mass_needed(1.0, 1000, 58.44) - 58.44) < 1e-3,     "mass_needed(1.0, 1000, 58.44) should be 58.44 g"

assert abs(percent_yield(3.87, 4.50) - (3.87/4.50*100)) < 0.01,     "percent_yield(3.87, 4.50) should be ≈ 86.0%"
assert abs(percent_yield(5.0, 5.0) - 100.0) < 1e-6,     "percent_yield(5.0, 5.0) should be 100.0%"

print("✓ All Practice 7 checks passed!")

```

---
## Section 8 — Reading and Writing Files

```python
with open("data.txt", "r") as fh:
    for line in fh:
        line = line.strip()   # removes the trailing newline
        ...
```


```python
# ── Example 8a — write a UV-Vis spectrum file ────────────────────────
spectrum_data = [
    (300, 0.021), (310, 0.045), (320, 0.098), (330, 0.187),
    (340, 0.312), (350, 0.458), (360, 0.581), (370, 0.634),
    (380, 0.589), (390, 0.492), (400, 0.371), (410, 0.261),
    (420, 0.163), (430, 0.095), (440, 0.052), (450, 0.028),
]
with open("uvvis_spectrum.txt", "w") as fh:
    fh.write("# Wavelength(nm), Absorbance\n")
    for wl, A in spectrum_data:
        fh.write(f"{wl}, {A:.3f}\n")
print("Written: uvvis_spectrum.txt")

```

```python
# ── Example 8b — read and find λ_max ────────────────────────────────
wavelengths, absorbances = [], []
with open("uvvis_spectrum.txt", "r") as fh:
    for line in fh:
        line = line.strip()
        if line.startswith("#") or not line:
            continue
        parts = line.split(", ")
        wavelengths.append(float(parts[0]))
        absorbances.append(float(parts[1]))

max_A      = max(absorbances)
lambda_max = wavelengths[absorbances.index(max_A)]
print(f"Read {len(wavelengths)} points.  λ_max = {lambda_max} nm  (A = {max_A:.3f})")

```

#### ✏️ Practice 8
`uvvis_spectrum.txt` is already on disk from the example above.

Fill in the three result variables:
- `wl_data`: list of wavelengths read from the file
- `abs_data`: list of absorbances read from the file
- `conc_uM`: list of concentrations in μM, using Beer–Lambert
  (`c = A / (ε·l)`, ε = 6400 L mol⁻¹ cm⁻¹, l = 1 cm, × 10⁶ for μM)

Then write a new file `uvvis_with_conc.txt` with three columns.


<!-- grade_id: practice8_solution, points: 0, grade: false, solution: true -->
```python
# ── Practice 8 ───────────────────────────────────────────────────────
# Given — do not change:
epsilon     = 6400.0   # L mol⁻¹ cm⁻¹
path_length = 1.0      # cm

# Results — fill these in:
wl_data  = None   # list of floats — wavelengths from file
abs_data = None   # list of floats — absorbances from file
conc_uM  = None   # list of floats — concentrations in μM

### BEGIN SOLUTION
wl_data  = []
abs_data = []
with open("uvvis_spectrum.txt", "r") as fh:
    for line in fh:
        line = line.strip()
        if line.startswith("#") or not line:
            continue
        parts = line.split(", ")
        wl_data.append(float(parts[0]))
        abs_data.append(float(parts[1]))

conc_uM = [A / (epsilon * path_length) * 1e6 for A in abs_data]

with open("uvvis_with_conc.txt", "w") as fh:
    fh.write("# Wavelength(nm), Absorbance, Concentration(uM)\n")
    for wl, A, c in zip(wl_data, abs_data, conc_uM):
        fh.write(f"{wl:.0f}, {A:.3f}, {c:.4f}\n")
### END SOLUTION

print(f"Read {len(wl_data)} data points.")
print("Written: uvvis_with_conc.txt")

```

<!-- grade_id: practice8_tests, points: 4, grade: true, solution: false -->
```python
# ── Tests for Practice 8 ─────────────────────────────────────────────
import os
assert len(wl_data)  == 16, f"wl_data should have 16 entries, got {len(wl_data)}"
assert len(abs_data) == 16, f"abs_data should have 16 entries, got {len(abs_data)}"
assert wl_data[0]  == 300.0, f"First wavelength should be 300 nm, got {wl_data[0]}"
assert wl_data[-1] == 450.0, f"Last wavelength should be 450 nm, got {wl_data[-1]}"

assert len(conc_uM) == 16, f"conc_uM should have 16 values, got {len(conc_uM)}"
expected_c0 = 0.021 / (6400.0 * 1.0) * 1e6
assert abs(conc_uM[0] - expected_c0) < 1e-6,     f"conc_uM[0] should be ≈ {expected_c0:.6f} μM, got {conc_uM[0]:.6f}"

assert os.path.exists("uvvis_with_conc.txt"),     "uvvis_with_conc.txt was not created — check your file-writing code"
with open("uvvis_with_conc.txt") as fh:
    data_lines = [l for l in fh if not l.startswith("#") and l.strip()]
assert len(data_lines) == 16,     f"uvvis_with_conc.txt should have 16 data lines, found {len(data_lines)}"

print("✓ All Practice 8 checks passed!")

```

---
## Capstone Exercise — Enzyme Kinetics

The **Michaelis–Menten equation**: $v_0 = V_{\max} \cdot [S] \;/\; (K_m + [S])$


```python
# ── Capstone setup — run this first ──────────────────────────────────
V_max = 120.0   # μmol/min
K_m   = 25.0    # μM
substrate_concs = [1, 2, 5, 10, 15, 20, 25, 30, 40, 50, 75, 100, 150, 200, 300]

def michaelis_menten(S, V_max, K_m):
    """Return reaction rate v0 for substrate concentration S."""
    return (V_max * S) / (K_m + S)

with open("kinetics_data.txt", "w") as fh:
    fh.write("# [S] (uM), v0 (umol/min)\n")
    for S in substrate_concs:
        fh.write(f"{S}, {michaelis_menten(S, V_max, K_m):.4f}\n")
print("Written: kinetics_data.txt  (V_max=120, Km=25)")

```

#### ✏️ Capstone — Your Turn
A second enzyme has **V_max = 85 μmol/min** and **K_m = 40 μM**.
The result variable names are pre-written — fill them in.

- `v0_new`: list of rates computed by calling `michaelis_menten` for each substrate concentration
- `S_new`: list of substrate concentrations read back from `kinetics_new.txt`
- `v0_new_read`: list of rates read back from `kinetics_new.txt`
- `km_estimate`: the substrate concentration (float) whose v₀ is closest to V_max_new / 2


<!-- grade_id: capstone_solution, points: 0, grade: false, solution: true -->
```python
# ── Capstone — Your Turn ─────────────────────────────────────────────
# Given — do not change:
V_max_new       = 85.0
K_m_new         = 40.0
substrate_concs = [1, 2, 5, 10, 15, 20, 25, 30, 40, 50, 75, 100, 150, 200, 300]

# Results — fill these in:
v0_new       = None   # list of floats — rates written to file
S_new        = None   # list of floats — [S] values read from file
v0_new_read  = None   # list of floats — rates read from file
km_estimate  = None   # float (μM) — [S] closest to half-maximal rate

### BEGIN SOLUTION
# 1. Write file
v0_new = []
with open("kinetics_new.txt", "w") as fh:
    fh.write("# [S] (uM), v0 (umol/min)\n")
    for S in substrate_concs:
        v0 = michaelis_menten(S, V_max_new, K_m_new)
        v0_new.append(v0)
        fh.write(f"{S}, {v0:.4f}\n")

# 2. Read file
S_new, v0_new_read = [], []
with open("kinetics_new.txt", "r") as fh:
    for line in fh:
        line = line.strip()
        if line.startswith("#") or not line:
            continue
        s_str, v_str = line.split(", ")
        S_new.append(float(s_str))
        v0_new_read.append(float(v_str))

# 3. Estimate Km
half_Vmax  = V_max_new / 2
diffs      = [abs(v - half_Vmax) for v in v0_new_read]
km_estimate = S_new[diffs.index(min(diffs))]
### END SOLUTION

print(f"Half-maximal rate closest at [S] = {km_estimate:.0f} μM  (true K_m = {K_m_new} μM)")

```

<!-- grade_id: capstone_tests, points: 4, grade: true, solution: false -->
```python
# ── Tests for Capstone ───────────────────────────────────────────────
import os
assert os.path.exists("kinetics_new.txt"), "kinetics_new.txt was not created"
assert len(v0_new) == 15, f"v0_new should have 15 values, got {len(v0_new)}"

expected_v0_40 = michaelis_menten(40, 85.0, 40.0)   # = 42.5
assert abs(v0_new[substrate_concs.index(40)] - expected_v0_40) < 1e-3,     f"v0 at [S]=40 μM should be ≈ {expected_v0_40:.2f}"

assert len(S_new) == 15 and len(v0_new_read) == 15,     "S_new and v0_new_read should each have 15 values"
assert km_estimate == 40.0,     f"km_estimate should be 40.0 μM (= K_m), got {km_estimate}"

print("✓ All Capstone checks passed!")

```

---
## Summary

| Concept | Key syntax | Chemistry use |
|---------|-----------|---------------|
| Variables | `x = 5.0` | Physical quantities |
| Arithmetic | `+`, `-`, `*`, `/`, `**` | Stoichiometry, unit conversion |
| Strings | `s.split()`, `f"..."`, `s.strip()` | Labels, formula parsing |
| Lists | `data[i]`, `append`, `len` | Spectral data, series of measurements |
| Loops | `for item in collection:` | Process all data points |
| Dictionaries | `d["key"]`, `d.items()` | Periodic table, spectral libraries |
| Functions | `def f(x): return ...` | Reusable equations |
| File I/O | `open()`, `readlines()` | Reading instrument output |

**Next notebook:** *Part 2 — Reading Gaussian Output Files*
